# Homework 1 (Total Points: 260) <a class="anchor" id="top"></a>


**Submission instructions**:
- Only the code `TODO: Implement this!` denotes that these sections are graded.
- For Part 1: You can use the `nltk`, `numpy` and `matplotlib` libraries here. Other libraries, e.g., `gensim` or `scikit-learn`, may not be used. For Part 2: `gensim` is allowed in addition to the imported libraries in the next code cell
- The notebook you submit has to have the student ids, separated by underscores (E.g., `12341234_12341234_12341234.ipynb`).
- This will be parsed by a regexp, so please double check your filename.
- Only one member of each group has to submit the file to canvas.
- Make sure to check that your notebook runs before submission. A quick way to do this is to restart the kernel and run all the cells.  
- Please do not delete/add new cells. Removing cells can lead to grade deduction.
- Note, that you are not allowed to use Google Colab.


**Learning Goals**:
- [Part 1, Term-based matching](#part1) (175 points):
    - Learn how to load a dataset and process it.
    - Learn how to implement several standard IR methods (TF-IDF, BM25, QL) and understand their weaknesses & strengths.
    - Learn how to evaluate IR methods.
- [Part 2, Semantic-based matching](#part2) (85 points):
    - Learn how to implement vector-space retrieval methods (LSI, LDA).
    - Learn how to use LSI and LDA for re-ranking

    
**Resources**:
- **Part 1**: Sections 2.3, 4.1, 4.2, 4.3, 5.3, 5.6, 5.7, 6.2, 7, 8 of [Search Engines: Information Retrieval in Practice](https://ciir.cs.umass.edu/downloads/SEIRiP.pdf)
- **Part 2**: [LSI - Chapter 18](https://nlp.stanford.edu/IR-book/pdf/18lsi.pdf) from [Introduction to Information Retrieval](https://nlp.stanford.edu/IR-book/) book and the [original LDA paper](https://jmlr.org/papers/volume3/blei03a/blei03a.pdf)

```python
# imports
# TODO: Ensure that no additional library is imported in the notebook.
# TODO: Only the standard library and the following libraries are allowed:
# TODO: You can also use unlisted classes from these libraries or standard libraries (such as defaultdict, Counter, ...).

import os
import zipfile
from functools import partial

import nltk
import requests
import numpy as np
from tqdm import tqdm

import matplotlib.pyplot as plt
from matplotlib.pyplot import cm

from ipywidgets import widgets
from IPython.display import display, HTML
#from IPython.html import widgets
from collections import namedtuple

%matplotlib inline
```


# Part 1: Term-based Matching (175 points) <a class="anchor" id="part1"></a>

[Back to top](#top)

In the first part, we will learn the basics of IR from loading and preprocessing the material, to implementing some well known search algorithms, to evaluating the ranking performance of the implemented algorithms. We will be using the CACM dataset throughout the assignment. The CACM dataset is a collection of titles and abstracts from the journal CACM (Communication of the ACM).

Table of contents:
- [Section 1: Text Processing](#text_processing) (15 points)
- [Section 2: Indexing](#indexing) (10 points)
- [Section 3: Ranking](#ranking) (80 points)
- [Section 4: Evaluation](#evaluation) (40 points)
- [Section 5: Analysis](#analysis) (30 points)


---
## Section 1: Text Processing (15 points)<a class="anchor" id="text_processing"></a>

[Back to Part 1](#part1)

In this section, we will load the dataset and learn how to clean up the data to make it usable for an IR system.
The points of this section are earned for the following implementations:
- `read_cacm_docs` (4 points): Reads in the CACM documents.
- `read_queries` (3 points): Reads in the CACM queries.
- `load_stopwords` (1 point): Loads the stopwords.
- `tokenize` (4 points): Tokenizes the input text.
- `stem_token` (3 points): Stems the given token.

We are using the [CACM dataset](http://ir.dcs.gla.ac.uk/resources/test_collections/cacm/), which is a small, classic IR dataset, composed of a collection of titles and abstracts from the journal CACM. It comes with relevance judgements for queries, so we can evaluate our IR system.


---
### 1.1 Read the CACM documents (4 points)


The following cell downloads the dataset and unzips it to a local directory.

```python
def download_dataset():
    folder_path = os.environ.get("IR1_DATA_PATH")
    if not folder_path:
        folder_path = "./datasets/"
    os.makedirs(folder_path, exist_ok=True)

    file_location = os.path.join(folder_path, "cacm.zip")

    # download file if it doesn't exist
    if not os.path.exists(file_location):

        url = "https://surfdrive.surf.nl/files/index.php/s/M0FGJpX2p8wDwxR/download"

        with open(file_location, "wb") as handle:
            print(f"Downloading file from {url} to {file_location}")
            response = requests.get(url, stream=True)
            for data in tqdm(response.iter_content()):
                handle.write(data)
            print("Finished downloading file")

    if not os.path.exists(os.path.join(folder_path, "train.txt")):

        # unzip file
        with zipfile.ZipFile(file_location, 'r') as zip_ref:
            zip_ref.extractall(folder_path)

download_dataset()
```

Output:
```
Downloading file from https://surfdrive.surf.nl/files/index.php/s/M0FGJpX2p8wDwxR/download to ./datasets/cacm.zip
706818it [00:05, 140265.60it/s]
Finished downloading file
```

---

You can see a brief description of each file in the dataset by looking at the README file:

```python
##### Read the README file
!cat ./datasets/README
#####
```

Output:
```
Files in this directory with sizes:
          0 Jun 19 21:01 README

    2187734 Jun 19 20:55 cacm.all              text of documents
        626 Jun 19 20:58 cite.info             key to citation info
                                                (the X sections in cacm.all)
       2668 Jun 19 20:55 common_words           stop words used by smart
       2194 Jun 19 20:55 make_coll*             shell script to make collection
       1557 Jun 19 20:55 make_coll_term*        ditto (both useless without
                                                smart system)
       9948 Jun 19 20:55 qrels.text             relation giving
                                                    qid did 0 0
                                                to indicate dument did is
                                                relevant to query qid
      13689 Jun 19 20:55 query.text             Original text of the query
```

---
We are interested in 4 files:
- `cacm.all` : Contains the text for all documents. Note that some documents do not have abstracts available.
- `query.text` : The text of all queries
- `qrels.text` : The relevance judgements
- `common_words` : A list of common words. This may be used as a collection of stopwords

```python
##### The first 45 lines of the CACM dataset forms the first record
# We are interested only in 3 fields.
# 1. the '.I' field, which is the document id
# 2. the '.T' field (the title) and
# 3. the '.W' field (the abstract, which may be absent)
!head -45 ./datasets/cacm.all
#####
```

Output:
```
.I 1
.T
Preliminary Report-International Algebraic Language
.B
CACM December, 1958
.A
Perlis, A. J.
Samelson,K.
.N
CA581203 JB March 22, 1978  8:28 PM
.X
100	5	1
123	5	1
164	5	1
1	5	1
1	5	1
1	5	1
205	5	1
210	5	1
214	5	1
1982	5	1
398	5	1
642	5	1
669	5	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
165	6	1
196	6	1
196	6	1
1273	6	1
1883	6	1
324	6	1
43	6	1
53	6	1
91	6	1
410	6	1
3184	6	1
```

```python
!head -200 ./datasets/cacm.all

```

Output:
```
.I 1
.T
Preliminary Report-International Algebraic Language
.B
CACM December, 1958
.A
Perlis, A. J.
Samelson,K.
.N
CA581203 JB March 22, 1978  8:28 PM
.X
100	5	1
123	5	1
164	5	1
1	5	1
1	5	1
1	5	1
205	5	1
210	5	1
214	5	1
1982	5	1
398	5	1
642	5	1
669	5	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
1	6	1
165	6	1
196	6	1
196	6	1
1273	6	1
1883	6	1
324	6	1
43	6	1
53	6	1
91	6	1
410	6	1
3184	6	1
.I 2
.T
Extraction of Roots by Repeated Subtractions for Digital Computers
.B
CACM December, 1958
.A
Sugai, I.
.N
CA581202 JB March 22, 1978  8:29 PM
.X
2	5	2
2	5	2
2	5	2
.I 3
.T
Techniques Department on Matrix Program Schemes
.B
CACM December, 1958
.A
Friedman, M. D.
.N
CA581201 JB March 22, 1978  8:30 PM
.X
3	5	3
3	5	3
3	5	3
.I 4
.T
Glossary of Computer Engineering and Programming Terminology
.B
CACM November, 1958
.N
CA581103 JB March 22, 1978  8:32 PM
.X
4	5	4
4	5	4
4	5	4
.I 5
.T
Two Square-Root Approximations
.B
CACM November, 1958
.A
Wadey, W. G.
.N
CA581102 JB March 22, 1978  8:33 PM
.
```

```python
!cat ./datasets/cite.info

```

Output:
```
4 : "bibliographic coupling" - if document id Y appears in the bibliographic
    coupling subvector for document X with a weight of w, it means X
    and Y have w common references in their bibliographies; the weight
    of did X in the vector for X is the number of items in X's bibliography.
5 : "links" - documents X and Y are linked if X cites Y, Y cites X, or
    X == Y.
6 : "co-citations" - if document id Y appears in the co-citation subvector
    for document X with weight w, it means X and Y are cited together in
    w documents; the weight of did X in the vector for X is the number
    of documents that cite X.
```

```python
%alias head powershell -command "& {Get-Content %s -Head 500}"
%head datasets/cacm.all
```

Output:
```
/bin/bash: powershell: command not found
```

---

**Implementation (4 points):**
Write a function to read the `cacm.all` file. Note that each document has a variable number of lines. The `.I` field denotes a new document:

<!-- grade_id: cell-b736116eb419c624, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (4 points)
def read_cacm_docs(root_folder = "./datasets/"):
    """
        Reads in the CACM documents. The dataset is assumed to be in the folder "./datasets/" by default
        Returns: A list of 2-tuples: (doc_id, document), where 'document' is a single string created by
            appending the title and abstract (separated by a "\n").
            In case the record doesn't have an abstract, the document is composed only by the title
    """
    list_tup = []
    ids = []
    txt = []
    with open(root_folder + "cacm.all", 'r+') as f:
        lines = f.readlines()

        for i in range(0, len(lines)):
            line = lines[i]

            if line[:2] == ".I":
                ids.append(line[3:-1])

            elif line[:2] == '.T':
                count = 2
                full_title = lines[i+1]
                title = lines[i+2]
                while title[:2] != '.B':

                    if title[:2] == '.W':
                        count += 1
                        full_title = full_title + lines[i+count]
                        count += 1
                        title = lines[i+count]

                        while title[:2] != '.B':
                            full_title = full_title[:-1] + ' ' + title
                            count += 1
                            title = lines[i+count]
                        break

                    full_title = full_title[:-1] + ' ' + title
                    count += 1
                    title = lines[i+count]

                txt.append(full_title[:-1])

    for i in range(len(ids)):
        list_tup.append((ids[i], txt[i]))

    return list_tup
```

<!-- grade_id: cell-a1c43818e0d3fd79, points: 4, grade: true, solution: false -->
```python
##### Function check
docs = read_cacm_docs()

assert isinstance(docs, list)
assert len(docs) == 3204, "There should be exactly 3024 documents"
#####
```

---
### 1.2 Read the CACM queries (3 points)

Next, let us read the queries. They are formatted similarly:

```python
##### The first 15 lines of 'query.text' has 2 queries
# We are interested only in 2 fields.
# 1. the '.I' - the query id
# 2. the '.W' - the query
!head -15 ./datasets/query.text
#####
```

Output:
```
.I 1
.W
 What articles exist which deal with TSS (Time Sharing System), an
operating system for IBM computers?
.N
 1. Richard Alexander, Comp Serv, Langmuir Lab (TSS)
 
.I 2
.W
 I am interested in articles written either by Prieve or Udo Pooch
.A
Prieve, B.
Pooch, U.
.N
 2. Richard Alexander, Comp Serv, Langmuir Lab (author = Pooch or Prieve)
```

```python
%head datasets/query.text
```

Output:
```
/bin/bash: powershell: command not found
```

---

**Implementation (3 points):**
Write a function to read the `query.text` file:

<!-- grade_id: cell-433e3ad5d0e2572a, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (3 points)
def read_queries(root_folder = "./datasets/"):
    """
        Reads in the CACM queries. The dataset is assumed to be in the folder "./datasets/" by default
        Returns: A list of 2-tuples: (query_id, query)
    """
    list_tup = []
    ids = []
    txt = []
    with open(root_folder + "query.text", 'r+') as f:
        lines = f.readlines()

        for i in range(0, len(lines)):
            line = lines[i]

            if line[:2] == ".I":
                ids.append(line[3:-1])

            elif line[:2] == '.W':
                count = 2
                full_title = lines[i+1][1:]
                title = lines[i+2]
                while title[:2] != '.N' and title[:2] != '.A':

                    full_title = full_title[:-1] + ' ' + title
                    count += 1
                    title = lines[i+count]

                txt.append(" ".join(full_title[:-1].split()))

    for i in range(len(ids)):
        list_tup.append((ids[i], txt[i]))

    return list_tup
```

<!-- grade_id: cell-6ec540abce66c598, points: 3, grade: true, solution: false -->
```python
##### Function check
queries = read_queries()

assert isinstance(queries, list)
assert len(queries) == 64 and all([q[1] is not None for q in queries]), "There should be exactly 64 queries"
#####
```

---
### 1.3 Read the stop words (1 point)

We use the common words stored in `common_words`:

```python
##### Read the stop words file
!head ./datasets/common_words
##### Read the README file
```

Output:
```
a
about
above
accordingly
across
after
afterwards
again
against
all
```

```python
%head datasets/common_words
```

Output:
```
/bin/bash: powershell: command not found
```

---
**Implementation (1 point):**
Write a function to read the `common_words` file (For better coverage, try to keep them in lowercase):

<!-- grade_id: cell-7357aa40f64e5bcb, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (1 point)
def load_stopwords(root_folder = "./datasets/"):
    """
        Loads the stopwords. The dataset is assumed to be in the folder "./datasets/" by default
        Output: A set of stopwords
    """
    set_stop = set()
    with open(root_folder + "common_words", 'r+') as f:
        lines = f.readlines()

        for i in range(0, len(lines)):
            line = lines[i]

            set_stop.add(line[:-1].lower())

    return set_stop
```

<!-- grade_id: cell-2ca3ac162004de97, points: 1, grade: true, solution: false -->
```python
##### Function check
stopwords = load_stopwords()

assert isinstance(stopwords, set)
assert len(stopwords) == 428, "There should be exactly 428 stop words"
#####
```

---
### 1.4 Tokenization (4 points)

We can now write some basic text processing functions.
A first step is to tokenize the text.

**Note**: Use the  `WordPunctTokenizer` available in the `nltk` library:

<!-- grade_id: cell-322be4c9499bdc4b, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (4 points)
def tokenize(text):
    """
        Tokenizes the input text. Use the WordPunctTokenizer
        Input: text - a string
        Output: a list of tokens
    """
    return nltk.WordPunctTokenizer().tokenize(text)
```

<!-- grade_id: cell-7fbf48bf7541a622, points: 4, grade: true, solution: false -->
```python
##### Function check
text = "the quick brown fox jumps over the lazy dog"
tokens = tokenize(text)

assert isinstance(tokens, list)
assert len(tokens) == 9

print(tokens)
# output: ['the', 'quick', 'brown', 'fox', 'jumps', 'over', 'the', 'lazy', 'dog']
#####
```

Output:
```
['the', 'quick', 'brown', 'fox', 'jumps', 'over', 'the', 'lazy', 'dog']
```

---
### 1.5 Stemming (3 points)

Write a function to stem tokens.
Again, you can use the nltk library for this:

<!-- grade_id: cell-e3f6c8e3f874b28d, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (3 points)
def stem_token(token):
    """
        Stems the given token using the PorterStemmer from the nltk library
        Input: a single token
        Output: the stem of the token
    """
    stemmer = nltk.PorterStemmer()
    return stemmer.stem(token)
```

<!-- grade_id: cell-cd6863e6ee6ed205, points: 3, grade: true, solution: false -->
```python
##### Function check

assert stem_token('owned') == 'own'
assert stem_token('itemization') == 'item'
#####
```

---
### 1.6 Summary

The following function puts it all together. Given a string, it tokenizes it and processes it according to the flags that you set.

```python
#### Putting it all together
def process_text(text, stem=False, remove_stopwords=False, lowercase_text=False):

    tokens = []
    for token in tokenize(text):
        if remove_stopwords and token.lower() in stopwords:
            continue
        if stem:
            token = stem_token(token)
        if lowercase_text:
            token = token.lower()
        tokens.append(token)

    return tokens
####
```

---

Let's create two sets of preprocessed documents.
We can process the documents and queries according to these two configurations:

```python
# In this configuration:
# Don't preprocess the text, except to tokenize
config_1 = {
  "stem": False,
  "remove_stopwords" : False,
  "lowercase_text": True
}


# In this configuration:
# Preprocess the text, stem and remove stopwords
config_2 = {
  "stem": True,
  "remove_stopwords" : True,
  "lowercase_text": True,
}

####
doc_repr_1 = []
doc_repr_2 = []
for (doc_id, document) in docs:
    doc_repr_1.append((doc_id, process_text(document, **config_1)))
    doc_repr_2.append((doc_id, process_text(document, **config_2)))

####
```

```python
doc_repr_2
```

Output:
```
[('1', ['preliminari', 'report', '-', 'intern', 'algebra', 'languag']),
 ('2', ['extract', 'root', 'repeat', 'subtract', 'digit', 'comput']),
 ('3', ['techniqu', 'depart', 'matrix', 'program', 'scheme']),
 ('4', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('5', ['squar', '-', 'root', 'approxim']),
 ('6', ['comput', 'inspect', 'procedur']),
 ('7', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('8', ['equival', 'transform', 'program', 'scheme']),
 ('9', ['propos', 'uncol']),
 ('10', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('11',
  ['problem',
   'program',
   'commun',
   'chang',
   'machin',
   'propos',
   'solut',
   '-',
   'part',
   '2']),
 ('12', ['error', 'estim', 'rung', '-', 'kutta', 'procedur']),
 ('13', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('14',
  ['problem',
   'program',
   'commun',
   'chang',
   'machin',
   'propos',
   'solut',
   '(',
   'part',
   '1',
   ')']),
 ('15', ['recurs', 'curv', 'fit
```

---

## Section 2: Indexing (10 points)<a class="anchor" id="indexing"></a>

[Back to Part 1](#part1)



A retrieval function usually takes in a query document pair, and scores a query against a document.  Our document set is quite small - just a few thousand documents. However, consider a web-scale dataset with a few million documents. In such a scenario, it would become infeasible to score every query and document pair. In such a case, we can build an inverted index. From Wikipedia:

> ... , an inverted index (also referred to as a postings file or inverted file) is a database index storing a mapping from content, such as words or numbers, to its locations in a table, .... The purpose of an inverted index is to allow fast full-text searches, at a cost of increased processing when a document is added to the database. ...


Consider a simple inverted index, which maps from word to document. This can improve the performance of a retrieval system significantly. In this assignment, we consider a *simple* inverted index, which maps a word to a set of documents. In practice, however, more complex indices might be used.  


### 2.1 Term Frequency-index (10 points)
In this assignment, we will be using an index created in memory since our dataset is tiny. To get started, build a simple index that maps each `token` to a list of `(doc_id, count)` where `count` is the count of the `token` in `doc_id`.
For consistency, build this index using a python dictionary.
    
Now, implement a function to build an index:

<!-- grade_id: cell-077599b87e953209, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (10 points)
def build_tf_index(documents):
    """
        Build an inverted index (with counts). The output is a dictionary which takes in a token
        and returns a list of (doc_id, count) where 'count' is the count of the 'token' in 'doc_id'
        Input: a list of documents - (doc_id, tokens)
        Output: An inverted index. [token] -> [(doc_id, token_count)]
    """
    #finding all unique tokens
    tokens_set = set()
    for doc in documents:
        for word in doc[1]:
            tokens_set.add(word)

    #creating a dict with counters for each doc
    inv_index = {}
    for token in tokens_set:
        inv_index[token] = []
        for doc in documents:
            counter = 0
            for word in doc[1]:
                if word == token:
                    counter += 1
            inv_index[token].append((doc[0], counter))

    return inv_index
```

```python
doc_repr_2
```

Output:
```
[('1', ['preliminari', 'report', '-', 'intern', 'algebra', 'languag']),
 ('2', ['extract', 'root', 'repeat', 'subtract', 'digit', 'comput']),
 ('3', ['techniqu', 'depart', 'matrix', 'program', 'scheme']),
 ('4', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('5', ['squar', '-', 'root', 'approxim']),
 ('6', ['comput', 'inspect', 'procedur']),
 ('7', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('8', ['equival', 'transform', 'program', 'scheme']),
 ('9', ['propos', 'uncol']),
 ('10', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('11',
  ['problem',
   'program',
   'commun',
   'chang',
   'machin',
   'propos',
   'solut',
   '-',
   'part',
   '2']),
 ('12', ['error', 'estim', 'rung', '-', 'kutta', 'procedur']),
 ('13', ['glossari', 'comput', 'engin', 'program', 'terminolog']),
 ('14',
  ['problem',
   'program',
   'commun',
   'chang',
   'machin',
   'propos',
   'solut',
   '(',
   'part',
   '1',
   ')']),
 ('15', ['recurs', 'curv', 'fit
```

---
Now we can build indexed documents and preprocess the queries based on the two configurations:

```python
#### Indexed documents based on the two configs

# Create the 2 indices
tf_index_1 = build_tf_index(doc_repr_1)
tf_index_2 = build_tf_index(doc_repr_2)

# This function returns the tf_index of the corresponding config
def get_index(index_set):
    assert index_set in {1, 2}
    return {
        1: tf_index_1,
        2: tf_index_2
    }[index_set]

####
#### Preprocessed query based on the two configs

# This function preprocesses the text given the index set, according to the specified config
def preprocess_query(text, index_set):
    assert index_set in {1, 2}
    if index_set == 1:
        return process_text(text, **config_1)
    elif index_set == 2:
        return process_text(text, **config_2)

####
```

<!-- grade_id: cell-fc7c7232d5d2ee46, points: 5, grade: true, solution: false -->
```python
##### Function check

print(tf_index_1['computer'][:10])
####
```

Output:
```
[('1', 0), ('2', 0), ('3', 0), ('4', 1), ('5', 0), ('6', 0), ('7', 1), ('8', 0), ('9', 0), ('10', 1)]
```

<!-- grade_id: cell-ff06bd11204db250, points: 5, grade: true, solution: false -->
```python
##### Function check

print(tf_index_2['computer'][:10])
####
```

Output:
```
[('1', 0), ('2', 0), ('3', 0), ('4', 0), ('5', 0), ('6', 0), ('7', 0), ('8', 0), ('9', 0), ('10', 0)]
```


---
## Section 3: Ranking  (80 points) <a class="anchor" id="ranking"></a>

[Back to Part 1](#part1)

Now that we have cleaned and processed our dataset, we can start building simple IR systems.

For now, we consider *simple* IR systems, which involve computing scores from the tokens present in the document/query. More advanced methods are covered in later assignments.

We will implement the following methods in this section:
- [Section 3.1: Bag of Words](#bow) (10 points)
- [Section 3.2: TF-IDF](#tfidf) (15 points)
- [Section 3.3: Query Likelihood Model](#qlm) (35 points)
- [Section 3.4: BM25](#bm25) (20 points)

**Scoring policy:**
Your implementations in this section are scored based on the expected performance of your ranking functions.
You will get a full mark if your implementation meets the expected performance (measured by some evaluation metric).
Otherwise, you may get partial credit.
For example, if your *Bag of words* ranking function has 60% of expected performance, you will get 4 out of 10.

---

### Section 3.1: Bag of Words (10 points)<a class="anchor" id="bow"></a>

Probably the simplest IR model is the Bag of Words (BOW) model.
Implement a function that scores and ranks all the documents against a query using this model.   

Note that you can use either the count of the token or 'binarize' it i.e set the value equal to 1 if the token appears.   


<!-- grade_id: cell-de9cf0459c4b9324, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (10 points)
def bow_search(query, index_set):
    """
        Perform a search over all documents with the given query.
        Note: You have to use the `get_index` function created in the previous cells
        Input:
            query - a (unprocessed) query
            index_set - the index to use
        Output: a list of (document_id, score), sorted in descending relevance to the given query
    """
    index = get_index(index_set)
    processed_query = preprocess_query(query, index_set)

    # YOUR CODE HERE
    #index is just an index I created before, dict with key as tokens and values as list with counts and ids
    #processed_query is the string of query after all preprocessing we chossed

    #creating a dictionary for all ids with scores for the query (initialized with 0s)
    dict_of_scores = {}
    for i_d in index[list(index)[0]]:
        dict_of_scores[i_d[0]] = float(0)

    #adding the scores to dict_of_scores (binary)
    for word in processed_query:
        if word in index:
            for i_d in index[word]:
                if i_d[1] != 0:
                    dict_of_scores[i_d[0]] += float(1)

    #translating a dictionary to list of tuples and sorting it
    list_of_scores = list(dict_of_scores.items())
    list_of_scores.sort(key=lambda tup: tup[1], reverse=True)

    return list_of_scores
```

```python

```

<!-- grade_id: cell-9f6aceae6dd9125f, points: 2, grade: true, solution: false -->
```python
#### Function check

docs_by_id = dict(docs)
def print_results(docs, len_limit=50):
    for i, (doc_id, score) in enumerate(docs):
        doc_content = docs_by_id[doc_id].strip().replace("\n", "\\n")[:len_limit] + "..."
        print(f"Rank {i}({score:.2}): {doc_content}")

test_bow = bow_search("report", index_set=1)[:5]
print(f"BOW Results:")
print_results(test_bow)
####
```

Output:
```
BOW Results:
Rank 0(1.0): Preliminary Report-International Algebraic Languag...
Rank 1(1.0): ALGOL Sub-Committee Report - Extensions...
Rank 2(1.0): The Use of Computers in Engineering Classroom Inst...
Rank 3(1.0): Report on a Conference of University Computing Cen...
Rank 4(1.0): Report on the Algorithmic Language ALGOL 60 .A Nau...
```

<!-- grade_id: cell-4eed3abf233d9b58, points: 2, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-4d65a2d7090c466c, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-dedf36ab5853ce20, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```


---

### Section 3.2: TF-IDF (15 points) <a class="anchor" id="tfidf"></a>

Before we implement the tf-idf scoring functions, let's first write a function to compute the document frequencies of all words.  

#### 3.2.1 Document frequency (5 points)
Compute the document frequencies of all tokens in the collection.  

<!-- grade_id: cell-9a2369f32e864b8a, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (5 points)
def compute_df(documents):
    """
        Compute the document frequency of all terms in the vocabulary
        Input: A list of documents
        Output: A dictionary with {token: document frequency)
    """
    # YOUR CODE HERE

    #finding all unique tokens
    tokens_set = set()
    for doc in documents:
        for word in doc:
            tokens_set.add(word)

    #creating a dict with counters for each doc
    dict_df = {}
    for token in tokens_set:
        dict_df[token] = 0
        for doc in documents:
            indicator = 0
            for word in doc:
                if word == token:
                    indicator = 1
            dict_df[token] += indicator

    return dict_df
```

```python
#### Compute df based on the two configs

# get the document frequencies of each document
df_1 = compute_df([d[1] for d in doc_repr_1])
df_2 = compute_df([d[1] for d in doc_repr_2])

def get_df(index_set):
    assert index_set in {1, 2}
    return {
        1: df_1,
        2: df_2
    }[index_set]
####
```

<!-- grade_id: cell-79e8a6db1e5fc46f, points: 5, grade: true, solution: false -->
```python
#### Function check

print(df_1['computer'])
print(df_2['computer'])
####
```

Output:
```
597
11
```

```python
index = get_index(index_set=1)
N = len(index[list(index)[0]])
N
```

Output:
```
3204
```

---
#### 3.2.2 TF-IDF search (10 points)
Next, implement a function that computes a tf-idf score given a query.      

<!-- grade_id: cell-2fb5ba34b2994cd9, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (10 points)
def tfidf_search(query, index_set):
    """
        Perform a search over all documents with the given query using tf-idf.
        Note #1: You have to use the `get_index` (and the `get_df`) function created in the previous cells
        Input:
            query - a (unprocessed) query
            index_set - the index to use
        Output: a list of (document_id, score), sorted in descending relevance to the given query
    """
    index = get_index(index_set)
    df = get_df(index_set)
    processed_query = preprocess_query(query, index_set)

    # YOUR CODE HERE
    #creating a dictionary for all ids with scores for the query (initialized with 0s)
    dict_of_scores = {}
    for i_d in index[list(index)[0]]:
        dict_of_scores[i_d[0]] = float(0)

    N = len(index[list(index)[0]])

    #adding scores for the query
    for word in processed_query:
        if word in index:
            for i_d in index[word]:
                if i_d[1] != 0:
                    dict_of_scores[i_d[0]] += i_d[1] * np.log(N/df[word])

    #translating a dictionary to list of tuples and sorting it
    list_of_scores = list(dict_of_scores.items())
    list_of_scores.sort(key=lambda tup: tup[1], reverse=True)

    return list_of_scores
```

<!-- grade_id: cell-bc68aeeacf42beb3, points: 2, grade: true, solution: false -->
```python
#### Function check
test_tfidf = tfidf_search("report", index_set=1)[:5]
print(f"TFIDF Results:")
print_results(test_tfidf)
####
```

Output:
```
TFIDF Results:
Rank 0(2.3e+01): An Information Algebra - Phase I Report-Language S...
Rank 1(2.3e+01): Rejuvenating Experimental Computer Science\nThis r...
Rank 2(1.2e+01): ALGOL 60 Confidential\nThe ALGOL 60 Report,* when ...
Rank 3(7.8): Automatic Abstracting and Indexing Survey and Reco...
Rank 4(7.8): A String Language for Symbol Manipulation Based on...
```

<!-- grade_id: cell-c7702fa8179fadb9, points: 2, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-3284f50ac29abbaa, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-d908c80a3155354b, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

---

### Section 3.3: Query Likelihood Model (35 points) <a class="anchor" id="qlm"></a>

In this section, you will implement a simple query likelihood model.


#### 3.3.1 Naive QL (15 points)

First, let us implement a naive version of a QL model, assuming a multinomial unigram language model (with a uniform prior over the documents).



```python
#### Document length for normalization

def doc_lengths(documents):
    doc_lengths = {doc_id:len(doc) for (doc_id, doc) in documents}
    return doc_lengths

doc_lengths_1 = doc_lengths(doc_repr_1)
doc_lengths_2 = doc_lengths(doc_repr_2)

def get_doc_lengths(index_set):
    assert index_set in {1, 2}
    return {
        1: doc_lengths_1,
        2: doc_lengths_2
    }[index_set]
####
```

<!-- grade_id: cell-8bcf2b804d636c2e, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (15 points)
def naive_ql_search(query, index_set):
    """
        Perform a search over all documents with the given query using a naive QL model.
        Note #1: You have to use the `get_index` (and get_doc_lengths) function created in the previous cells
        Input:
            query - a (unprocessed) query
            index_set - the index to use
        Output: a list of (document_id, score), sorted in descending relevance to the given query
    """
    index = get_index(index_set)
    doc_lengths = get_doc_lengths(index_set)
    processed_query = preprocess_query(query, index_set)
    # YOUR CODE HERE

    dict_of_scores = {}
    for i_d in index[list(index)[0]]:
        dict_of_scores[i_d[0]] = float(1) ## how to initialize??

    for word in processed_query:
        if word in index:
            print ("word",word)
            for docID_count in index[word]: # index[word] is a list: get value (list) of the key (word). i_d is each element in the list. i_d is (doc_id, counter)
                docID = docID_count[0]
                dict_of_scores[docID] = dict_of_scores[docID] * (docID_count[1]/doc_lengths[docID])



    #translating a dictionary to list of tuples and sorting it by descending order
    list_of_scores = list(dict_of_scores.items())
    list_of_scores.sort(key=lambda tup: tup[1], reverse=True)
    print ("list_of_scores",list_of_scores)
    return list_of_scores
#naive_ql_search("report", index_set=1)
   #raise NotImplementedError()
```

<!-- grade_id: cell-5a83ac12ecde8578, points: 3, grade: true, solution: false -->
```python
#### Function check
test_naiveql = naive_ql_search("report", index_set=1)[:5]
print(f"Naive QL Results:")
print_results(test_naiveql)
####
```

Output:
```
word report
list_of_scores [('599', 0.2), ('2689', 0.2), ('1', 0.16666666666666666), ('947', 0.16666666666666666), ('65', 0.14285714285714285), ('584', 0.14285714285714285), ('985', 0.125), ('147', 0.1111111111111111), ('601', 0.1111111111111111), ('984', 0.08333333333333333), ('2583', 0.08333333333333333), ('689', 0.06666666666666667), ('196', 0.058823529411764705), ('946', 0.045454545454545456), ('1086', 0.041666666666666664), ('3184', 0.03773584905660377), ('1416', 0.03125), ('1531', 0.02857142857142857), ('3160', 0.02553191489361702), ('616', 0.020134228187919462), ('721', 0.019230769230769232), ('691', 0.017543859649122806), ('329', 0.016666666666666666), ('1659', 0.015873015873015872), ('2479', 0.015503875968992248), ('1927', 0.015151515151515152), ('644', 0.015037593984962405), ('2226', 0.014492753623188406), ('2930', 0.014492753623188406), ('321', 0.013043478260869565), ('2048', 0.01282051282051282), ('146', 0.012345679012345678), ('1765', 0.012195121951219513), ('1771', 0.0109
```

<!-- grade_id: cell-80f4bf2137f997bb, points: 2, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-5ce2993458a8ce51, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-7753bdb54e292f3d, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-54e476e2f96e64bb, points: 4, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

---
#### 3.3.2 QL (20 points)
Now, let's implement a QL model that handles the issues with the naive version. In particular, you will implement a QL model with Jelinek-Mercer Smoothing. That means an interpolated score is computed per word - one term is the same as the previous naive version, and the second term comes from a unigram language model. In addition, you should accumulate the scores by summing the **log** (smoothed) probability which leads to better numerical stability.

<!-- grade_id: cell-bb1f506409771257, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (20 points)

# YOUR CODE HERE
#raise NotImplementedError()

def ql_search(query, index_set):
    """
        Perform a search over all documents with the given query using a QL model
        with Jelinek-Mercer Smoothing (set smoothing=0.1).


        Note #1: You have to use the `get_index` (and get_doc_lengths) function created in the previous cells
        Note #2: You might have to create some variables beforehand and use them in this function


        Input:
            query - a (unprocessed) query
            index_set - the index to use
        Output: a list of (document_id, score), sorted in descending relevance to the given query
    """
    index = get_index(index_set)
    doc_lengths = get_doc_lengths(index_set)
    processed_query = preprocess_query(query, index_set)

    # YOUR CODE HERE
    smoothing = 0.1
    none_smoothing =1-smoothing
    collection_len = len (index)
    print ("collection_len", collection_len)
    dict_of_scores = {}
    for i_d in index[list(index)[0]]:
        dict_of_scores[i_d[0]] = float(1)

    for word in processed_query:
        if word in index:
            l = index[word]
            freq_in_collection = sum([pair[1] for pair in l])
            print ("word",word, "freq_in_collection", freq_in_collection)
            for docID_count in index[word]: # index[word] is a list. eg. index[word]=tf_index_1['computer']= value (a list) of the key (word). i_d is each element in the list. i_d is (doc_id, counter)
                docID = docID_count[0]
                smoothed_prob = none_smoothing * (docID_count[1]/doc_lengths[docID]) + smoothing * (freq_in_collection/collection_len)
                dict_of_scores[docID] = dict_of_scores[docID] * smoothed_prob


#     print ("dict_of_scores",dict_of_scores)
    list_of_scores = list(dict_of_scores.items())
    list_of_scores.sort(key=lambda tup: tup[1], reverse=True)
#     print ("list_of_scores",list_of_scores)
    return list_of_scores

```

<!-- grade_id: cell-850e9d6369bcec32, points: 4, grade: true, solution: false -->
```python
#### Function check
test_ql_results = ql_search("report", index_set=1)[:5]
print_results(test_ql_results)
print()
test_ql_results_long = ql_search("report " * 10, index_set=1)[:5]
print_results(test_ql_results_long)
####
```

Output:
```
collection_len 9698
word report freq_in_collection 84
Rank 0(0.18): A Report Writer For COBOL...
Rank 1(0.18): A CRT Report Generating System...
Rank 2(0.15): Preliminary Report-International Algebraic Languag...
Rank 3(0.15): Supplement to the ALGOL 60 Report...
Rank 4(0.13): ALGOL Sub-Committee Report - Extensions...

collection_len 9698
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
word report freq_in_collection 84
Rank 0(3.7e-08): A Report Writer For COBOL...
Rank 1(3.7e-08): A CRT Report Generating System...
Rank 2(6.1e-09): Preliminary Report-International Algebraic Languag...
Rank 3(6.1e-09): Supplement to the ALGOL 60 Report...
Rank 4(1.3e-09): ALGOL Sub-Committee Report - Extensions...
```

<!-- grade_id: cell-958cdcf6fd6899b7, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-384dc23a0c251f6e, points: 4, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-7218966cba5097cc, points: 4, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-481ab073259ae53f, points: 5, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

---

### Section 3.4: BM25 (20 points) <a class="anchor" id="bm25"></a>

In this section, we will implement the BM25 scoring function.


<!-- grade_id: cell-15640fc9b5d00a3c, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (20 points)
def bm25_search(query, index_set):
    """
        Perform a search over all documents with the given query using BM25. Use k_1 = 1.5 and b = 0.75
        Note #1: You have to use the `get_index` (and `get_doc_lengths`) function created in the previous cells
        Note #2: You might have to create some variables beforehand and use them in this function

        Input:
            query - a (unprocessed) query
            index_set - the index to use
        Output: a list of (document_id, score), sorted in descending relevance to the given query
    """

    k=1.5
    b = 0.75
    index = get_index(index_set)
    df = get_df(index_set)
    doc_lengths = get_doc_lengths(index_set)
    processed_query = preprocess_query(query, index_set)

    #list with (doc_id,score)
    bmRes = []
    avg_doc_length = sum(list(doc_lengths.values()))/len(list(doc_lengths.values()))

    for i,doc in enumerate(docs):

        bmResult = 0
        #we loop through every word in query and compute bm25 score for every score
        for word in processed_query:
            if word in index:
                tf_w = [tf for tf in index[word] if tf[0]==doc[0]][0][1]  #term frequency calculation
                idf = np.log(len(index[word]) / df[word])   #inverse document freq calculation
                bmResult += idf*((tf_w*(k+1))/(tf_w + k*(1-b+b*doc_lengths[doc[0]]/avg_doc_length)))#bm25 score formula

        bmRes.append((doc[0],bmResult))

    bmRes = sorted(bmRes,key=lambda tup: tup[1], reverse=True) #sort list by score

    return bmRes





```

```python
docs_by_id = dict(docs)
def print_results(docs, len_limit=50):
    for i, (doc_id, score) in enumerate(docs):
        #if doc_id in ['2434','2863','3078']:
            doc_content = docs_by_id[doc_id].strip().replace("\n", "\\n")[:len_limit] + "..."
            print(f"Rank {i} {doc_id}({score:.2}): {doc_content}")

```

```python
test_bm25_results = bm25_search(queries[20][1], index_set=1)[:10]
print_results(test_bm25_results)
print('~~~~~~~~~~~~~~~~~~~~~~~~~~~~~~')
print_results(bow_search(queries[20][1], index_set=1))[:10]
```

Output:
```
Rank 0 2325(1.9e+01): Numerical Mathematics and Computer Science\nNumeri...
Rank 1 1309(1.4e+01): A Computer User-Oriented System\nA computer langua...
Rank 2 1619(1.4e+01): Error-Free Methods for Statistical Computations\nN...
Rank 3 2936(1.3e+01): An Efficient Data Structure for the Simulation Eve...
Rank 4 3166(1.3e+01): Computing Standard Deviations: Accuracy\nFour algo...
Rank 5 2702(1.3e+01): On the Complexity of LR(k) Testing\nThe problem of...
Rank 6 3018(1.2e+01): Covering Edges by Cliques with Regard to Keyword C...
Rank 7 3163(1.2e+01): An Optimal Insertion Algorithm for One-Sided Heigh...
Rank 8 2950(1.2e+01): A Unifying Approach to Scheduling\nThis paper pres...
Rank 9 2226(1.2e+01): Further Evidence for the Analysis of Algorithms fo...
~~~~~~~~~~~~~~~~~~~~~~~~~~~~~~
Rank 0 2189(8.0): Generation of Rosary Permutations Expressed in Ham...
Rank 1 2931(8.0): Logic and Programming Languages\nLogic has been lo...
Rank 2 719(7.0): Variable Width Stacks\nCharacter addressable, va
```

```python
queries[1][1]
```

<!-- grade_id: cell-d10536bca72c74b1, points: 3, grade: true, solution: false -->
```python
#### Function check
test_bm25_results = bm25_search("report", index_set=1)[:5]
print_results(test_bm25_results)
####
```

```python
#### Function check
bm25_search("report", index_set=1)
#print_results(test_bm25_results)
####
```

<!-- grade_id: cell-60f6ec5052712d79, points: 5, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-5d17524043a5abcc, points: 5, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-ff8e704eda1184e3, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

<!-- grade_id: cell-a52310500a2543cb, points: 4, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```


---

### 3.5. Test Your Functions

The widget below allows you to play with the search functions you've written so far. Use this to test your search functions and ensure that they work as expected.

```python
#### Highlighter function
# class for results
ResultRow = namedtuple("ResultRow", ["doc_id", "snippet", "score"])
# doc_id -> doc
docs_by_id = dict((d[0], d[1]) for d in docs)

def highlight_text(document, query, tol=17):
    import re
    tokens = tokenize(query)
    regex = "|".join(f"(\\b{t}\\b)" for t in tokens)
    regex = re.compile(regex, flags=re.IGNORECASE)
    output = ""
    i = 0
    for m in regex.finditer(document):
        start_idx = max(0, m.start() - tol)
        end_idx = min(len(document), m.end() + tol)
        output += "".join(["...",
                        document[start_idx:m.start()],
                        "<strong>",
                        document[m.start():m.end()],
                        "</strong>",
                        document[m.end():end_idx],
                        "..."])
    return output.replace("\n", " ")


def make_results(query, search_fn, index_set):
    results = []
    for doc_id, score in search_fn(query, index_set):
        highlight = highlight_text(docs_by_id[doc_id], query)
        if len(highlight.strip()) == 0:
            highlight = docs_by_id[doc_id]
        results.append(ResultRow(doc_id, highlight, score))
    return results
####
```

```python
# TODO: Set this to the function you want to test!
# this function should take in a query (string)
# and return a sorted list of (doc_id, score)
# with the most relevant document in the first position
search_fn = bm25_search
index_set = 1

text = widgets.Text(description="Search Bar", width=200)
display(text)

def handle_submit(sender):
    print(f"Searching for: '{sender.value}'")

    results = make_results(sender.value, search_fn, index_set)

    # display only the top 5
    results = results[:10]

    body = ""
    for idx, r in enumerate(results):
        body += f"<li>Document #{r.doc_id}({r.score}): {r.snippet}</li>"
    display(HTML(f"<ul>{body}</ul>"))


text.on_submit(handle_submit)
```

---

## Section 4: Evaluation (40 points) <a class="anchor" id="evaluation"></a>

[Back to Part 1](#part1)

Before we jump in and implement an algorithm for retrieval, we first have to learn how to evaluate such a system. In particular, we will work with offline evaluation metrics. These metrics are computed on a dataset with known relevance judgements.

Implement the following evaluation metrics.

1. Precision (7 points)
2. Recall (7 points)
3. Mean Average Precision (12 points)
4. Expected Reciprocal Rank (12 points)

---
### 4.1 Read relevance labels (2 points)

Let's take a look at the `qrels.text` file, which contains the ground truth relevance scores. The relevance labels for CACM are binary - either 0 or 1.


```python
!head ./datasets/qrels.text
```

```python
!head datasets/qrels.text
```

---
**Implementation (2 points):**
The first column is the query_id and the second column is the document_id. You can safely ignore the 3rd and 4th columns. Write a function to read in the file:

<!-- grade_id: cell-ee5253a4ef602fce, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (2 points)
def read_qrels(root_folder = "./datasets/"):
    """
        Reads the qrels.text file.
        Output: A dictionary: query_id -> [list of relevant documents]
    """
    # YOUR CODE HERE
    dict_qrels = {}
    with open(root_folder + "qrels.text", 'r+') as f:
        lines = f.readlines()

        for i in range(0, len(lines)):
            line = lines[i]
            if line[0] == '0':
                qid = line[1:2]
            else:
                qid = line[0:2]

            if qid not in dict_qrels:
                    dict_qrels[qid] = [line[3:7]]
            else:
                dict_qrels[qid].append(line[3:7])

    return dict_qrels
```

<!-- grade_id: cell-72215605fbe24f65, points: 2, grade: true, solution: false -->
```python
#### Function check
qrels = read_qrels()

assert len(qrels) == 52, "There should be 52 queries with relevance judgements"
assert sum(len(j) for j in qrels.values()) == 796, "There should be a total of 796 Relevance Judgements"
####
```

---
**Note:** For a given query `query_id`, you can assume that documents *not* in `qrels[query_id]` are not relevant to `query_id`.


```python
print(qrels)
```

---
### 4.2 Precision (7 points)
Implement the `precision@k` metric:

<!-- grade_id: cell-494bd0cce108ed67, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (7 points)
def precision_k(results, relevant_docs, k):
    """
        Compute Precision@K
        Input:
            results: A sorted list of 2-tuples (document_id, score),
                    with the most relevant document in the first position
            relevant_docs: A set of relevant documents.
            k: the cut-off
        Output: Precision@K
    """
    counter = 0
    for doc in results[:k]:
        if doc[0] in relevant_docs:
            counter += 1
    print(counter, "",k)
    return counter/k
```

<!-- grade_id: cell-e7ff0d91c319ca64, points: 7, grade: true, solution: false -->
```python
#### Function check
qid = queries[0][0]
qtext = queries[0][1]
print(f'query:{qtext}')
results = bm25_search(qtext, 2)
precision = precision_k(results, qrels[qid], 10)
print(f'precision@10 = {precision}')
####
```

```python
qrels[qid]
```

```python
qid = queries[10][0]
qtext = queries[10][1]
print(f'query:{qtext}')
results = bm25_search(qtext, 2)
precision = precision_k(results,qrels[qid],10)
print(f'precision@10 = {precision}')
```

---
### 4.3 Recall (7 points)
Implement the `recall@k` metric:

<!-- grade_id: cell-c323fc8c3f8a7cf8, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (7 points)
def recall_k(results, relevant_docs, k):
    """
        Compute Recall@K
        Input:
            results: A sorted list of 2-tuples (document_id, score), with the most relevant document in the first position
            relevant_docs: A set of relevant documents.
            k: the cut-off
        Output: Recall@K
    """
    # YOUR CODE HERE
    counter = 0
    for doc in results[:k]:
        if doc[0] in relevant_docs:
            counter += 1
    #print(counter," ",len(relevant_docs))
    return counter/len(relevant_docs)
```

```python
docs[0]
```

<!-- grade_id: cell-b25172161aef165c, points: 7, grade: true, solution: false -->
```python
#### Function check
qid = queries[10][0]
qtext = queries[10][1]
print(f'query:{qtext}')
results = bm25_search(qtext, 2)
recall = recall_k(results, qrels[qid], 10)
print(f'recall@10 = {recall}')
####
```

---
### 4.4 Mean Average Precision (12 points)
Implement the `map` metric:

<!-- grade_id: cell-e50925fa9093a30d, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (12 points)
def average_precision(results, relevant_docs):
    """
        Compute Average Precision (for a single query - the results are
        averaged across queries to get MAP in the next few cells)
        Hint: You can use the recall_k and precision_k functions here!
        Input:
            results: A sorted list of 2-tuples (document_id, score), with the most
                    relevant document in the first position
            relevant_docs: A set of relevant documents.
        Output: Average Precision
    """
    # YOUR CODE HERE
    count = 0
    for i in range(1,(len(docs)+1)):
        if results[i-1][0] in relevant_docs:
             count += precision_k(results, relevant_docs, i)
    return count / len(relevant_docs)




    #raise NotImplementedError()

```

<!-- grade_id: cell-8a1f7ec98571e58b, points: 12, grade: true, solution: false -->
```python
#### Function check
qid = queries[20][0]
qtext = queries[20][1]
print(f'query:{qtext}')
results = bm25_search(qtext, 2)
mean_ap = average_precision(results, qrels[qid])
print(f'MAP = {mean_ap}')
####
```

```python
#### Function check
qid = queries[0][0]
qtext = queries[0][1]
print(f'query:{qtext}')
results = bm25_search(qtext, 2)
mean_ap = average_precision(results, qrels[qid])
print(f'MAP = {mean_ap}')
####
```

---
### 4.5 Expected Reciprocal Rank (12 points)
Implement the `err` metric:

<!-- grade_id: cell-64262889f9b267ea, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (12 points)
def err(results, relevant_docs):
    """
        Compute the expected reciprocal rank.
        Hint: https://dl.acm.org/doi/pdf/10.1145/1645953.1646033?download=true
        Input:
            results: A sorted list of 2-tuples (document_id, score), with the most
                    relevant document in the first position
            relevant_docs: A set of relevant documents.
        Output: ERR

    """

    count = 0
    for i in results:
        count += 1
        if i[0] in relevant_docs:
            print(i,count)
            break
    return 1 / count


```

```python
results[0]
#relevant_docs[0]
```

<!-- grade_id: cell-071e3970ff1afae4, points: 12, grade: true, solution: false -->
```python
#### Function check
qid = queries[30][0]
qtext = queries[30][1]
print(f'query:{qtext}')
results = bm25_search(qtext, 2)
ERR = err(results, qrels[qid])
print(f'ERR = {ERR}')
####
```

```python
qid = queries[1][0]
qtext = queries[1][1]
print(f'query:{qtext}')
results = bm25_search(qtext, 2)
ERR = err(results, qrels[qid])
print(f'ERR = {ERR}')
```

---
### 4.6 Evaluate Search Functions

Let's define some metrics@k using [partial functions](https://docs.python.org/3/library/functools.html#functools.partial)

```python
#### metrics@k functions

recall_at_1 = partial(recall_k, k=1)
recall_at_5 = partial(recall_k, k=5)
recall_at_10 = partial(recall_k, k=10)
precision_at_1 = partial(precision_k, k=1)
precision_at_5 = partial(precision_k, k=5)
precision_at_10 = partial(precision_k, k=10)


list_of_metrics = [
    ("ERR", err),
    ("MAP", average_precision),
    ("Recall@1",recall_at_1),
    ("Recall@5", recall_at_5),
    ("Recall@10", recall_at_10),
    ("Precision@1", precision_at_1),
    ("Precision@5", precision_at_5),
    ("Precision@10", precision_at_10)]
####
```

---

The following function evaluates a `search_fn` using the `metric_fn`. Note that the final number is averaged over all the queries

```python
#### Evaluate a search function

list_of_search_fns = [
    ("BOW", bow_search),
    ("TF-IDF", tfidf_search),
    ("NaiveQL", naive_ql_search),
    ("QL", ql_search),
    ("BM25", bm25_search)
]

def evaluate_search_fn(search_fn, metric_fns, index_set=None):
    # build a dict query_id -> query
    queries_by_id = dict((q[0], q[1]) for q in queries)

    metrics = {}
    for metric, metric_fn in metric_fns:
        metrics[metric] = np.zeros(len(qrels), dtype=np.float32)

    for i, (query_id, relevant_docs) in enumerate(qrels.items()):
        query = queries_by_id[query_id]
        if index_set:
            results = search_fn(query, index_set)
        else:
            results = search_fn(query)

        for metric, metric_fn in metric_fns:
            metrics[metric][i] = metric_fn(results, relevant_docs)



    final_dict = {}
    for metric, metric_vals in metrics.items():
        final_dict[metric] = metric_vals.mean()

    return final_dict
####
```

## Section 5: Analysis (30 points) <a class="anchor" id="analysis"></a>

[Back to Part 1](#part1)

In the final section of Part1, we will compare the different term-based IR algorithms and different preprocessing configurations and analyze their advantages and disadvantages.

### Section 5.1: Plot (20 points)

First, gather the results. The results should consider the index set, the different search functions and different metrics. Plot the results in bar charts, per metric, with clear labels.

**Rubric:**
- Each Metric is plotted: 7 points
- Each Method is plotted: 7 points
- Clear titles, x label, y labels and legends (if applicable): 6 points

<!-- grade_id: cell-46fda42a25863a04, points: 20, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-71-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

---
### Section 5.2: Summary (10 points)
Write a summary of what you observe in the results.
Your summary should compare results across the 2 indices and the methods being used. State what you expected to see in the results, followed by either supporting evidence *or* justify why the results did not support your expectations.      

Write your answer here!

---
---
# Part 2: Semantic-based Matching (85 points) <a class="anchor" id="part2"></a>

[Back to top](#top)

We will now experiment with methods that go beyond lexical methods like TF-IDF, which operate at the word level and are high dimensional and sparse, and look at methods which constructs low dimensional dense representations of queries and documents.

Since these low-dimensional methods have a higher time complexity, they are typically used in conjunction with methods like BM-25. That is, instead of searching through potentially million documents to find matches using low dimensional vectors, a list of K documents are retrieved using BM25, and then **re-ranked** using the other method. This is the method that is going to be applied in the following exercises.

LSI/LDA takes documents that are similar on a semantic level - for instance, if they are describing the same topic - and projects them into nearby vectors, despite having low lexical overlap.

In this assignment, you will use `gensim` to create LSI/LDA models and use them in re-ranking.

**Note**: The following exercises only uses `doc_repr_2` and `config_2`

Table of contents:
- [Section 6: LSI](#lsi) (15 points)
- [Section 7: LDA](#lda) (10 points)
- [Section 8: Word2Vec/Doc2Vec](#2vec) (20 points)
- [Section 8: Re-ranking](#reranking) (10 points)
- [Section 9: Re-ranking Evaluation](#reranking_eval) (30 points)

---
## Section 6: Latent Semantic Indexing (LSI) (15 points) <a class="anchor" id="lsi"></a>

[Back to Part 2](#part2)

LSI is one of the methods to embed the queries and documents into vectors. It is based on a method similar to Principal Component Analysis (PCA) for obtaining a dense concept matrix out of the sparse term-document matrix.

See [wikipedia](https://en.wikipedia.org/wiki/Latent_semantic_analysis), particularly [#Mathematics_of_LSI](https://en.wikipedia.org/wiki/Latent_semantic_analysis#Mathematics_of_LSI).

```python
import gensim
```

```python
from gensim.corpora import Dictionary
from gensim.models import LdaModel, LsiModel, Word2Vec
from gensim.models.doc2vec import Doc2Vec, TaggedDocument
from gensim import downloader as g_downloader
# gensim uses logging, so set it up
import logging
logging.basicConfig(format='%(asctime)s : %(levelname)s : %(message)s', level=logging.INFO)
```

---
### Section 6.1: Cosine Similarity (5 points)<a class="anchor" id="cosing_sim"></a>
Before we begin, let us first define our method of similarity for the LSI model, the cosine similarity:

$$\text{similarity} = \cos(\theta) = {\mathbf{A} \cdot \mathbf{B} \over \|\mathbf{A}\| \|\mathbf{B}\|} = \frac{ \sum\limits_{i=1}^{n}{A_i  B_i} }{ \sqrt{\sum\limits_{i=1}^{n}{A_i^2}}  \sqrt{\sum\limits_{i=1}^{n}{B_i^2}} }$$

Since we are using gensim, the types of vectors returned by their classes are of the form defined below (they are not just simple vectors):

```python
# 1, 2, 3 are either latent dimensions (LSI), or topics (LDA)
# The second value in each tuple is a number (LSI) or a probability (LDA)
example_vec_1 = [(1, 0.2), (2, 0.3), (3, 0.4)]
example_vec_2 = [(1, 0.2), (2, 0.7), (3, 0.4)]
```

---
**Implementation (2+3 points):**
Now, implement the `dot product` operation on these types of vectors and using this operator, implement the `cosine similarity` (don't forget: two functions to implement!):

<!-- grade_id: cell-0e8189f5f93de33f, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (2 points)
# TODO: Implement this! (2 points)
import math
def dot(vec_1,vec_2):
    """
        vec_1 and vec_2 are of the form: [(int, float), (int, float), ...]
        Return the dot product of two such vectors, computed only on the floats
        You can assume that the lengths of the vectors are the same, and the dimensions are aligned
            i.e you won't get: vec_1 = [(1, 0.2)] ; vec_2 = [(2, 0.3)]
                                (dimensions are unaligned and lengths are different)
    """
    # YOUR CODE HERE
    result=0
    for num1, num2 in zip(vec_1, vec_2):
        result = result+num1[1]* num2[1]
    return result



# TODO: Implement this! (3 points)
def cosine_sim(vec_1, vec_2):
    denominator = math.sqrt(dot(vec_1, vec_1)) * math.sqrt(dot(vec_2,vec_2))
    if denominator==0:
        cosine = 0
    else:
        cosine = dot(vec_1,vec_2)/denominator
    return cosine
```

<!-- grade_id: cell-b25d04ed6b79fd35, points: 2, grade: true, solution: false -->
```python
##### Function check
print(f'vectors: {(example_vec_1,example_vec_2)}')
print(f'dot product = {dot(example_vec_1,example_vec_2)}')
print(f'cosine similarity = {cosine_sim(example_vec_1,example_vec_2)}')
#####
```

Output:
```
vectors: ([(1, 0.2), (2, 0.3), (3, 0.4)], [(1, 0.2), (2, 0.7), (3, 0.4)])
dot product = 0.41000000000000003
cosine similarity = 0.9165587597202866
```

<!-- grade_id: cell-ae3c4466866ace77, points: 3, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

---
### Section 6.2: LSI Retrieval (10 points)<a class="anchor" id="lsi_retrieval"></a>
LSI retrieval is simply ranking the documents based on their cosine similarity to the query vector.
First, let's write a parent class for vector-based retrieval models:

```python
class VectorSpaceRetrievalModel:
    """
        Parent class for Dense Vector Retrieval models
    """
    def __init__(self, doc_repr):
        """
            document_collection:
                [
                    (doc_id_1, [token 1, token 2, ...]),
                    (doc_id_2, [token 1, token 2, ....])
                    ...
                ]

        """
        self.doc_repr = doc_repr
        self.documents = [_[1] for _ in self.doc_repr]

        # construct a dictionary
        self.dictionary = Dictionary(self.documents)
        # Filter out words that occur less than 20 documents, or more than 50% of the documents.
        self.dictionary.filter_extremes(no_below=10)
        self.corpus = [self.dictionary.doc2bow(doc) for doc in self.documents]

        # Make a index to word dictionary.
        temp = self.dictionary[0]  # This is only to "load" the dictionary.
        self.id2word = self.dictionary.id2token

        # this is set by the train_model function
        self.model = None


    def vectorize_documents(self):
        """
            Returns a doc_id -> vector dictionary
        """
        vectors = {}
        for (doc_id, _), cc in zip(self.doc_repr, self.corpus):
            vectors[doc_id] = self.model[cc]
        return vectors

    def vectorize_query(self, query):
        # Note the use of config_2 here!
        query = process_text(query, **config_2)
        query_vector = self.dictionary.doc2bow(query)
        return self.model[query_vector]

    def train_model(self):
        """
            Trains a model and sets the 'self.model' variable.
            Make sure to use the variables created in the __init__ method.
            e.g the variables which may be useful: {corpus, dictionary, id2word}
        """
        raise NotImplementedError()
```

```python
x = VectorSpaceRetrievalModel(doc_repr_2)
#print(x.doc_repr)
v = x.vectorize_documents()
print(v)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-19-0c6ba5ce27f0> in <module>()
----> 1 x = VectorSpaceRetrievalModel(doc_repr_2)
      2 #print(x.doc_repr)
      3 v = x.vectorize_documents()
      4 print(v)

NameError: name 'doc_repr_2' is not defined
```

---
**Implementation (5 points):**
Implement the `train_model` method in the following class (note that this is only one line of code in `gensim`!). Ensure that the parameters defined in the `__init__` method are not changed, and are *used in the `train_method` function*. Normally, the hyperaparameter space will be searched using grid search / other methods - in this assignment we have provided the hyperparameters for you.

The last two lines of code train an LSI model on the list of documents which have been stemmed, lower-cased and have stopwords removed.

<!-- grade_id: cell-307682c9089f15d6, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (5 points)
class LsiRetrievalModel(VectorSpaceRetrievalModel):

    def __init__(self, doc_repr):
        super().__init__(doc_repr)

        self.num_topics = 100
        self.chunksize = 2000

    def train_model(self):
        # YOUR CODE HERE
        self.model = LsiModel(self.corpus,
                              num_topics=self.num_topics,chunksize = self.chunksize)


```

```python
#l.train_model()
doc_repr_2
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-20-25a06928c364> in <module>()
      1 #l.train_model()
----> 2 doc_repr_2

NameError: name 'doc_repr_2' is not defined
```

```python
l = LsiRetrievalModel(doc_repr_2)

```

<!-- grade_id: cell-5ce512650c1b2dfb, points: 0, grade: true, solution: false -->
```python
##### Function check
lsi = LsiRetrievalModel(doc_repr_2)
lsi.train_model()

# you can now get an LSI vector for a given query in the following way:
lsi.vectorize_query("report")
#####
```

```python
##### Function check
lsi = LsiRetrievalModel(doc_repr_2)
lsi.train_model()

# you can now get an LSI vector for a given query in the following way:
lsi.vectorize_documents()
```

\#### Please do not change this. This cell is used for grading.

---
Next, implement a basic ranking class for vector space retrieval (used for all semantic methods):

<!-- grade_id: cell-250515d288e80cdc, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (5 points)
class DenseRetrievalRanker:
    def __init__(self, vsrm, similarity_fn):
        """
            vsrm: instance of class `VectorSpaceRetrievalModel`
            similarity_fn: function instance that takes in two vectors
                            and returns a similarity score e.g cosine_sim defined earlier
        """
        self.vsrm = vsrm
        self.vectorized_documents = self.vsrm.vectorize_documents()
        self.similarity_fn = similarity_fn


    def _compute_sim(self, query_vector):
        """
            Compute the similarity of `query_vector` to documents in
            `self.vectorized_documents` using `self.similarity_fn`
            Returns a list of (doc_id, score) tuples
        """
        dict_of_scores = {}
        for key in self.vectorized_documents:
            dict_of_scores[key] = self.similarity_fn (self.vectorized_documents[key],
                                                      query_vector) # query_vector: a list, len=100
        list_of_scores = list(dict_of_scores.items())
        return list_of_scores

    def search(self, query):
        scores = self._compute_sim(self.vsrm.vectorize_query(query))
        scores.sort(key=lambda _:-_[1])
        return scores
```

<!-- grade_id: cell-b73068b3e77a8e31, points: 0, grade: true, solution: false -->
```python
##### Function check
drm_lsi = DenseRetrievalRanker(lsi, cosine_sim)
drm_lsi.search("report")[:5]
#####
```

\#### Please do not change this. This cell is used for grading.

---
Now, you can test your LSI model in the following cell: try finding queries which are lexically different to documents, but semantically similar - does LSI work well for these queries?!

```python
# test your LSI model
search_fn = drm_lsi.search

text = widgets.Text(description="Search Bar", width=200)
display(text)

def make_results_2(query, search_fn):
    results = []
    for doc_id, score in search_fn(query):
        highlight = highlight_text(docs_by_id[doc_id], query)
        if len(highlight.strip()) == 0:
            highlight = docs_by_id[doc_id]
        results.append(ResultRow(doc_id, highlight, score))
    return results

def handle_submit_2(sender):
    print(f"Searching for: '{sender.value}' (SEARCH FN: {search_fn})")

    results = make_results_2(sender.value, search_fn)

    # display only the top 5
    results = results[:5]

    body = ""
    for idx, r in enumerate(results):
        body += f"<li>Document #{r.doc_id}({r.score}): {r.snippet}</li>"
    display(HTML(f"<ul>{body}</ul>"))


text.on_submit(handle_submit_2)
```

---
## Section 7: Latent Dirichlet Allocation (LDA) (10 points) <a class="anchor" id="lda"></a>

[Back to Part 2](#part2)

The specifics of LDA is out of the scope of this assignment, but we will use the `gensim` implementation to perform search using LDA over our small document collection. The key thing to remember is that LDA, unlike LSI, outputs a topic **distribution**, not a vector. With that in mind, let's first define a similarity measure.


---
### Section 7.1: Jenson-Shannon divergence (5 points) <a class="anchor" id="js_sim"></a>

The Jenson-Shannon divergence is a symmetric and finite measure on two probability distributions (unlike the KL, which is neither). For identical distributions, the JSD is equal to 0, and since our code uses 0 as irrelevant and higher scores as relevant, we use `(1 - JSD)` as the score or 'similarity' in our setup

**Note**: the JSD is bounded to \[0,1\] only if we use log base 2. So please ensure that you're using `np.log2` instead of `np.log`

<!-- grade_id: cell-d2376a85a4841e98, points: 0, grade: false, solution: true -->
```python
## TODO: Implement this! (5 points)
def jenson_shannon_divergence(vec_1, vec_2, assert_prob=False):
    """
        Computes the Jensen-Shannon divergence between two probability distributions.
        NOTE: DO NOT RETURN 1 - JSD here, that is handled by the next function which is already implemented!
        The inputs are *gensim* vectors - same as the vectors for the cosine_sim function
        assert_prob is a flag that checks if the inputs are proper probability distributions
            i.e they sum to 1 and are positive - use this to check your inputs if needed.
                (This is optional to implement, but recommended -
                you can the default to False to save a few ms off the runtime)
    """
    # YOUR CODE HERE
    if assert_prob == True:
        sum_1 = 0
        sum_2 = 0
        for prob in vec_1:
            if prob[1] < 0:
                print('Negative probability, wrong vector vec_1 as input.')
            sum_1 += prob[1]

        for prob in vec_2:
            if prob[1] < 0:
                print('Negative probability, wrong vector vec_2 as input.')
            sum_2 += prob[1]

        if sum_1 != 1:
            print('Probability distribiution of vec_1 does not sum up to 1!')
        if sum_2 != 1:
            print('Probability distribiutions of vec_2 does not sum up to 1!')

    M = []
    for i in range(len(vec_1)):
        M.append((vec_1[i][0], 1/2*(vec_1[i][1]+vec_2[i][1])))

    vec_1_array = np.asarray(vec_1, dtype=np.float)
    vec_2_array = np.asarray(vec_2, dtype=np.float)
    M_array = np.asarray(M, dtype=np.float)

    D_vec_1_M = np.sum(np.where(vec_1_array != 0, vec_1_array * np.log2(vec_1_array / M_array), 0))
    D_vec_2_M = np.sum(np.where(vec_2_array != 0, vec_2_array * np.log2(vec_2_array / M_array), 0))

    return 1/2*D_vec_1_M + 1/2*D_vec_2_M

def jenson_shannon_sim(vec_1, vec_2, assert_prob=False):
    return 1 - jenson_shannon_divergence(vec_1, vec_2)

```

<!-- grade_id: cell-487c6d2933f38053, points: 5, grade: true, solution: false -->
```python
##### Function check
vec_1 = [(1, 0.3), (2, 0.4), (3, 0.3)]
vec_2 = [(1, 0.1), (2, 0.7), (3, 0.2)]
jenson_shannon_sim(vec_1, vec_2, assert_prob=True)
#####
```

Output:
```
0.9251064410358459
```

---
### Section 7.2: LDA retrieval (5 points) <a class="anchor" id="lda_ret"></a>

Implement the `train_model` method in the following class (note that this is only one line of code in `gensim`!). Ensure that the parameters defined in the `__init__` method are not changed, and are *used in the `train_method` function*. You do not need to set this. Normally, the hyperaparameter space will be searched using grid search / other methods. Note that training the LDA model might take some time

The last two lines of code train an LDA model on the list of documents which have been stemmed, lower-cased and have stopwords removed.

<!-- grade_id: cell-021a48dff4a8bb91, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (5 points)
class LdaRetrievalModel(VectorSpaceRetrievalModel):
    def __init__(self, doc_repr):
        super().__init__(doc_repr)

        # use these parameters in the train_model method
        self.num_topics = 100
        self.chunksize = 2000
        self.passes = 20
        self.iterations = 400
        self.eval_every = 10
        # this is need to get full vectors
        self.minimum_probability=0.0
        self.alpha='auto'
        self.eta='auto'


    def train_model(self):
        # YOUR CODE HERE
        self.model = LdaModel(corpus = self.corpus, num_topics = self.num_topics, id2word = self.id2word,
                              chunksize = self.chunksize, passes = self.passes, iterations = self.iterations,
                              eval_every = self.eval_every, minimum_probability = self.minimum_probability,
                              alpha = self.alpha, eta = self.eta)
```

<!-- grade_id: cell-86750b715f0345fd, points: 0, grade: true, solution: false -->
```python
##### Function check
lda = LdaRetrievalModel(doc_repr_2)
lda.train_model()

# you can now get an LDA vector for a given query in the following way:
lda.vectorize_query("report")
#####
```

Output:
```
2021-02-19 14:42:55,419 : INFO : adding document #0 to Dictionary(0 unique tokens: [])
2021-02-19 14:42:55,550 : INFO : built Dictionary(5940 unique tokens: ['-', 'algebra', 'intern', 'languag', 'preliminari']...) from 3204 documents (total 115982 corpus positions)
2021-02-19 14:42:55,561 : INFO : discarding 4743 tokens: [('repeat', 8), ('glossari', 7), ('inspect', 8), ('uncol', 2), ('rung', 9), ('secant', 2), ('.', 1605), ('acceler', 6), ('diverg', 3), ('induc', 9)]...
2021-02-19 14:42:55,562 : INFO : keeping 1197 tokens which were in no less than 10 and no more than 1602 (=50.0%) documents
2021-02-19 14:42:55,566 : INFO : resulting dictionary: Dictionary(1197 unique tokens: ['-', 'algebra', 'intern', 'languag', 'preliminari']...)
2021-02-19 14:42:55,646 : INFO : using autotuned alpha, starting with [0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0.01, 0
```

\#### Please do not change this. This cell is used for grading.

---
Now we can use the `DenseRetrievalModel` class to obtain an LDA search function.
You can test your LDA model in the following cell: Try finding queries which are lexically different to documents, but semantically similar - does LDA work well for these queries?!

```python
drm_lda = DenseRetrievalRanker(lda, jenson_shannon_sim)

# test your LDA model
search_fn = drm_lda.search

text = widgets.Text(description="Search Bar", width=200)
display(text)


text.on_submit(handle_submit_2)
```

Output:
```
Text(value='', description='Search Bar')
```

## Section 8: Word2Vec/Doc2Vec (20 points) <a class="anchor" id="2vec"></a>

[Back to Part 2](#part2)

We will implement two other methods here, the Word2Vec model and the Doc2Vec model, also using `gensim`. Word2Vec creates representations of words, not documents, so the word level vectors need to be aggregated to obtain a representation for the document. Here, we will simply take the mean of the vectors.


A drawback of these models is that they need a lot of training data. Our dataset is tiny, so in addition to using a model trained on the data, we will also use a pre-trained model for Word2Vec (this will be automatically downloaded).     

*Note*:
1. The code in vectorize_documents / vectorize_query should return gensim-like vectors i.e `[(dim, val), .. (dim, val)]`.
2. For Word2Vec: You should also handle the following two cases: (a) A word in the query is not present in the vocabulary of the model and (b) none of the words in the query are present in the model - you can return 0 scores for all documents in this case. For either of these, you can check if a `word` is present in the vocab by using `word in self.model`


<!-- grade_id: cell-2b73759f9baf688f, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (10 points)
import pandas as pd
class W2VRetrievalModel(VectorSpaceRetrievalModel):
    my_model = None
    def __init__(self, doc_repr):
        super().__init__(doc_repr)
        self.my_corpus = [x[1] for x in doc_repr]

        # the dimensionality of the vectors
        self.size = 100
        self.min_count = 1


    def train_model(self):
        """
        Trains the W2V model
        """
        # YOUR CODE HERE
        self.my_model = Word2Vec(self.my_corpus,size=100, min_count=1)
        self.my_model.save("word2vec.model")
        #raise NotImplementedError()

    def vectorize_documents(self):
        """
            Returns a doc_id -> vector dictionary
        """
        # YOUR CODE HERE
        #doc_vector = model.wv
        word_list={}

        for d in self.doc_repr:
            gensim_list=[]
            new_list=[v for v in d[1] if v in self.my_model]
            if not new_list: continue


            result = [sum(x) for x in zip(*self.my_model[new_list])]
            results = [x/len(new_list) for x in result]
            for i,r in enumerate(results):
                gensim_list.append((i+1,r))
        #gensim_list.append((i,(sum(w)/len(w))))
            word_list[d[0]] = gensim_list
        return word_list
        #raise NotImplementedError()


    def vectorize_query(self, query):
        """
        Vectorizes the query using the W2V model
        """
        query = process_text(query, **config_2)

        # YOUR CODE HERE
        q_new_list = [v for v in query if v in self.my_model]

        result = [sum(x) for x in zip(*self.my_model[q_new_list])]
        results = [x/len(q_new_list) for x in result]
#         print(self.my_model[q_new_list])
#         print("---------------")
#         print(results)
#         print("---------------")
#         print(result)
        return results


        #raise NotImplementedError()


class W2VPretrainedRetrievalModel(W2VRetrievalModel):
    def __init__(self, doc_repr):
        super().__init__(doc_repr)
        self.model_name = "word2vec-google-news-300"
        self.size = 300

    def train_model(self):
        """
        Loads the pretrained model
        """
        self.my_model = g_downloader.load(self.model_name)

w2v = W2VRetrievalModel(doc_repr_2)
w2v.train_model()

# you can now get a W2V vector for a given query in the following way:
w2v.vectorize_query("report")
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-4717c64b4d57> in <module>()
      1 # TODO: Implement this! (10 points)
      2 import pandas as pd
----> 3 class W2VRetrievalModel(VectorSpaceRetrievalModel):
      4     my_model = None
      5     def __init__(self, doc_repr):

NameError: name 'VectorSpaceRetrievalModel' is not defined
```

```python
w = W2VRetrievalModel(doc_repr_2)
w.train_model()

print(w.vectorize_documents()["1"])
```

Output:
```
2021-02-19 14:49:39,480 : INFO : adding document #0 to Dictionary(0 unique tokens: [])
2021-02-19 14:49:39,598 : INFO : built Dictionary(5940 unique tokens: ['-', 'algebra', 'intern', 'languag', 'preliminari']...) from 3204 documents (total 115982 corpus positions)
2021-02-19 14:49:39,607 : INFO : discarding 4743 tokens: [('repeat', 8), ('glossari', 7), ('inspect', 8), ('uncol', 2), ('rung', 9), ('secant', 2), ('.', 1605), ('acceler', 6), ('diverg', 3), ('induc', 9)]...
2021-02-19 14:49:39,608 : INFO : keeping 1197 tokens which were in no less than 10 and no more than 1602 (=50.0%) documents
2021-02-19 14:49:39,613 : INFO : resulting dictionary: Dictionary(1197 unique tokens: ['-', 'algebra', 'intern', 'languag', 'preliminari']...)
2021-02-19 14:49:40,693 : INFO : collecting all words and their counts
2021-02-19 14:49:40,694 : INFO : PROGRESS: at sentence #0, processed 0 words, keeping 0 word types
2021-02-19 14:49:40,719 : INFO : collected 5940 word types from a corpus of 115982 raw w
```

<!-- grade_id: cell-b31c0f8d214b8bdf, points: 0, grade: true, solution: false -->
```python
assert len(w2v.vectorize_query("report")) == 100
assert len(w2v.vectorize_query("this is a sentence that is not mellifluous")) == 100

```

Output:
```
/usr/local/lib/python3.6/dist-packages/ipykernel_launcher.py:54: DeprecationWarning: Call to deprecated `__contains__` (Method will be removed in 4.0.0, use self.wv.__contains__() instead).
/usr/local/lib/python3.6/dist-packages/ipykernel_launcher.py:56: DeprecationWarning: Call to deprecated `__getitem__` (Method will be removed in 4.0.0, use self.wv.__getitem__() instead).
```

\#### Please do not change this. This cell is used for grading.

```python
w2v_pretrained = W2VPretrainedRetrievalModel(doc_repr_2)
w2v_pretrained.train_model()

# you can now get an W2V vector for a given query in the following way:
w2v_pretrained.vectorize_query("report")
```

Output:
```
2021-02-19 14:49:52,895 : INFO : adding document #0 to Dictionary(0 unique tokens: [])
2021-02-19 14:49:53,019 : INFO : built Dictionary(5940 unique tokens: ['-', 'algebra', 'intern', 'languag', 'preliminari']...) from 3204 documents (total 115982 corpus positions)
2021-02-19 14:49:53,029 : INFO : discarding 4743 tokens: [('repeat', 8), ('glossari', 7), ('inspect', 8), ('uncol', 2), ('rung', 9), ('secant', 2), ('.', 1605), ('acceler', 6), ('diverg', 3), ('induc', 9)]...
2021-02-19 14:49:53,030 : INFO : keeping 1197 tokens which were in no less than 10 and no more than 1602 (=50.0%) documents
2021-02-19 14:49:53,035 : INFO : resulting dictionary: Dictionary(1197 unique tokens: ['-', 'algebra', 'intern', 'languag', 'preliminari']...)
2021-02-19 14:49:53,110 : INFO : Creating /root/gensim-data
[==================================================] 100.0% 1662.8/1662.8MB downloaded
2021-02-19 14:54:05,656 : INFO : word2vec-google-news-300 downloaded
2021-02-19 14:54:05,674 : INFO : loading p
```

<!-- grade_id: cell-1b1466f8ce516f42, points: 2, grade: true, solution: false -->
```python
##### Function check

print(len(w2v_pretrained.vectorize_query("report")))
#####
```

Output:
```
<ipython-input-101-54d628cab034> in vectorize_query(self, query)
     52 
     53         # YOUR CODE HERE
---> 54         q_new_list = [v for v in query if v in self.my_model]
     55 
     56         result = [sum(x) for x in zip(*self.my_model[q_new_list])]

<ipython-input-101-54d628cab034> in <listcomp>(.0)
     52 
     53         # YOUR CODE HERE
---> 54         q_new_list = [v for v in query if v in self.my_model]
     55 
     56         result = [sum(x) for x in zip(*self.my_model[q_new_list])]

TypeError: argument of type 'NoneType' is not iterable
```

```python
drm_w2v = DenseRetrievalRanker(w2v, cosine_sim)

# test your LDA model
search_fn = drm_w2v.search

text = widgets.Text(description="Search Bar", width=200)
display(text)


text.on_submit(handle_submit_2)
```

Output:
```
/usr/local/lib/python3.6/dist-packages/ipykernel_launcher.py:33: DeprecationWarning: Call to deprecated `__contains__` (Method will be removed in 4.0.0, use self.wv.__contains__() instead).
/usr/local/lib/python3.6/dist-packages/ipykernel_launcher.py:37: DeprecationWarning: Call to deprecated `__getitem__` (Method will be removed in 4.0.0, use self.wv.__getitem__() instead).
Text(value='', description='Search Bar')
Searching for: 'report' (SEARCH FN: <bound method DenseRetrievalRanker.search of <__main__.DenseRetrievalRanker object at 0x7f15565b9358>>)
/usr/local/lib/python3.6/dist-packages/ipykernel_launcher.py:54: DeprecationWarning: Call to deprecated `__contains__` (Method will be removed in 4.0.0, use self.wv.__contains__() instead).
/usr/local/lib/python3.6/dist-packages/ipykernel_launcher.py:56: DeprecationWarning: Call to deprecated `__getitem__` (Method will be removed in 4.0.0, use self.wv.__getitem__() instead).
<ipython-input-75-8ae4997b2264> in cosine_sim(vec_1, vec_2)
   
```

```python
drm_w2v_pretrained = DenseRetrievalRanker(w2v_pretrained, cosine_sim)

# test your LDA model
search_fn = drm_w2v_pretrained.search

text = widgets.Text(description="Search Bar", width=200)
display(text)


text.on_submit(handle_submit_2)
```

Output:
```
<ipython-input-101-54d628cab034> in vectorize_documents(self)
     31         for d in self.doc_repr:
     32             gensim_list=[]
---> 33             new_list=[v for v in d[1] if v in self.my_model]
     34             if not new_list: continue
     35 

<ipython-input-101-54d628cab034> in <listcomp>(.0)
     31         for d in self.doc_repr:
     32             gensim_list=[]
---> 33             new_list=[v for v in d[1] if v in self.my_model]
     34             if not new_list: continue
     35 

TypeError: argument of type 'NoneType' is not iterable
```

**Implementation (10 points):**
For Doc2Vec, you will need to create a list of `TaggedDocument` instead of using the `self.corpus` or `self.documents` variable. Use the document id as the 'tag'.
  

<!-- grade_id: cell-680facdcc98a19ab, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (10 points)
class D2VRetrievalModel(VectorSpaceRetrievalModel):
    def __init__(self, doc_repr):
        super().__init__(doc_repr)

        self.vector_size= 100
        self.min_count = 1
        self.epochs = 20

        # YOUR CODE HERE
        documents = [TaggedDocument(doc, [i]), doc in enumerate(doc_repr)]
        raise NotImplementedError()

    def train_model(self):
        # YOUR CODE HERE
        model = Doc2Vec(documents, self.vector_size=100, self.min_count, self.epochs)
        raise NotImplementedError()

    def vectorize_documents(self):
        """
            Returns a doc_id -> vector dictionary
        """
        # YOUR CODE HERE
        vector =  model.infer_vector
        raise NotImplementedError()

    def vectorize_query(self, query):
        # YOUR CODE HERE
        raise NotImplementedError()

d2v = D2VRetrievalModel(doc_repr_2)
d2v.train_model()


# # you can now get an LSI vector for a given query in the following way:
d2v.vectorize_query("report")
```

<!-- grade_id: cell-5e2c5e0c9a2e8cb5, points: 0, grade: true, solution: false -->
```python
#### Please do not change this. This cell is used for grading.
```

\#### Please do not change this. This cell is used for grading.

```python
drm_d2v = DenseRetrievalRanker(d2v, cosine_sim)

# test your LDA model
search_fn = drm_d2v.search

text = widgets.Text(description="Search Bar", width=200)
display(text)


text.on_submit(handle_submit_2)
```

---
## Section 9: Re-ranking (10 points) <a class="anchor" id="reranking"></a>

[Back to Part 2](#part2)

To motivate the re-ranking perspective (i.e retrieve with lexical method + rerank with a semantic method), let's search using semantic methods and compare it to BM25's performance, along with their runtime:


```python
query = "algebraic functions"
print("BM25: ")
%timeit bm25_search(query, 2)
print("LSI: ")
%timeit drm_lsi.search(query)
print("LDA: ")
%timeit drm_lda.search(query)
print("W2V: ")
%timeit drm_w2v.search(query)
print("W2V(Pretrained): ")
%timeit drm_w2v_pretrained.search(query)
print("D2V:")
%timeit drm_d2v.search(query)
```

---

**Implementation (10 points):**
Re-ranking involves retrieving a small set of documents using simple but fast methods like BM25 and then re-ranking them with the aid of semantic methods such as LDA or LSI. Implement the following class, which takes in an `initial_retrieval_fn` - the initial retrieval function and `vsrm` - an instance of the `VectorSpaceRetrievalModel` class (i.e LSI/LDA) as input. The search function should first retrieve an initial list of K documents, and then these documents are re-ranked using a semantic method. This not only makes retrieval faster, but semantic methods perform poorly when used in isolation, as you will find out.

<!-- grade_id: cell-5bf47600d1a0c507, points: 0, grade: false, solution: true -->
```python
# TODO: Implement this! (10 points)
class DenseRerankingModel:
    def __init__(self, initial_retrieval_fn, vsrm, similarity_fn):
        """
            initial_retrieval_fn: takes in a query and returns a list of [(doc_id, score)] (sorted)
            vsrm: instance of `VectorSpaceRetrievalModel`
            similarity_fn: function instance that takes in two vectors
                            and returns a similarity score e.g cosine_sim defined earlier
        """
        self.ret = initial_retrieval_fn
        self.vsrm = vsrm
        self.similarity_fn = similarity_fn
        self.vectorized_documents = vsrm.vectorize_documents()

        assert len(self.vectorized_documents) == len(doc_repr_2)

    def search(self, query, K=50):
        """
            First, retrieve the top K results using the retrieval function
            Then, re-rank the results using the VSRM instance
        """
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-52c6d18a4c0b4882, points: 0, grade: true, solution: false -->
```python
##### Function check
bm25_search_2 = partial(bm25_search, index_set=2)
lsi_rerank = DenseRerankingModel(bm25_search_2, lsi, cosine_sim)
lda_rerank = DenseRerankingModel(bm25_search_2, lda, jenson_shannon_sim)
w2v_rerank = DenseRerankingModel(bm25_search_2, w2v, cosine_sim)
w2v_pretrained_rerank = DenseRerankingModel(bm25_search_2, w2v_pretrained, cosine_sim)
d2v_rerank = DenseRerankingModel(bm25_search_2, d2v, cosine_sim)

#####
```

\#### Please do not change this. This cell is used for grading.

---
Now, let us time the new search functions:

```python
query = "algebraic functions"
print("BM25: ")
%timeit bm25_search(query, 2)
print("LSI: ")
%timeit lsi_rerank.search(query)
print("LDA: ")
%timeit lda_rerank.search(query)
print("W2V: ")
%timeit w2v_rerank.search(query)
print("W2V(Pretrained): ")
%timeit w2v_pretrained_rerank.search(query)
print("D2V:")
%timeit d2v_rerank.search(query)
```

---
As you can see, it is much faster (but BM25 is still orders of magnitude faster).

---
## Section 10: Evaluation & Analysis (30 points) <a class="anchor" id="reranking_eval"></a>

[Back to Part 2](#part2)

[Previously](#evaluation) we have implemented some evaluation metrics and used them for measuring the ranking performance of term-based IR algorithms. In this section, we will do the same for semantic methods, both with and without re-ranking.

### Section 10.1: Plot (10 points)

First, gather the results. The results should consider the index set, the different search functions and different metrics. Plot the results in bar charts, per metric, with clear labels.

Then, gather only the re-ranking models, and plot and compare them with the results obtained in part 1 (only index set 2).

```python
list_of_sem_search_fns = [
    ("lda", drm_lda.search),
    ("lsi", drm_lsi.search),
    ("w2v", drm_w2v.search),
    ("w2v_pretrained", drm_w2v_pretrained.search),
    ("d2v", drm_d2v.search),
    ("lsi_rr", lsi_rerank.search),
    ("lda_rr", lda_rerank.search),
    ("w2v_rr", w2v_rerank.search),
    ("w2v_pretrained_rr", w2v_pretrained_rerank.search),
    ("d2v_rr", d2v_rerank.search),

]
```

<!-- grade_id: cell-7dd8273b0f5a3c22, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Section 10.2: Summary (20 points)

Your summary should compare methods from Part 1 and Part 2 (only for index set 2). State what you expected to see in the results, followed by either supporting evidence *or* justify why the results did not support your expectations. Consider the availability of data, scalability, domain/type of data, etc.

<!-- grade_id: cell-ec5dd7d9cf59dd86, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE